```python
# Initialize Otter
import otter
grader = otter.Notebook("lab09.ipynb")
```

<img src="data6.png" style="width: 15%; float: right; padding: 1%; margin-right: 2%;"/>

# Lab 09 – Functions, Control, and Iteration

## Data 6, Summer 2026

Welcome to Lab 09! In this lab, we'll be covering **functions**, **control**, and **iteration**.

```python
# Just run this cell
from datascience import *
import numpy as np
import warnings
warnings.simplefilter('ignore')

import seaborn as sns
import matplotlib.pyplot as plt
plt.style.use("fivethirtyeight")
%matplotlib inline
```

<hr style="border: 1px solid #fdb515;" />

# [Tutorial] Part 1: Boolean Expressions

---

## [Tutorial] Comparison Operators

In Python, the boolean is a data type with only two possible values: `True` and `False`. Expressions containing comparison operators such as `<` (less than), `>` (greater than), and `==` (equal to) evaluate to Boolean values. A list of common comparison operators:

| Comparison	| Operator	| True Example	| False Example | 
| --- | --- | --- | --- |
| Less than | 	<	| 2 < 3 	 |2 < 2 | 
| Greater than | 	>	 | 3 > 2	 | 3 > 3 | 
| Less than or equal | 	<= | 	2 <= 2 | 	3 <= 2 | 
| Greater or equal | 	>= | 	3 >= 3 | 	2 >= 3 | 
| Equal | 	==	 | 3 == 3 | 	3 == 2 | 
| Not equal | 	!=	 | 3 != 2	 | 2 != 2 | 

Run the following series of cells to see some examples of these comparison operators in practice. For each cell, make sure you understand the cell output before continuing.

```python
1 < 2
```

```python
min(20, 30) <= max(-10, 20)
```

```python
make_array(1, 2, 3, 4) != 3
```

```python
"abc" < "def"
```

Note that == and = are **not** the same!
    
>`=` is used for **assignment**, i.e., to assign a Python name to a value.<br/>On the other hand, `==` is used as the **comparison operator** to compare two values. 

```python
result = (3 + 4) == (3 * 4)
result
```

```python
# Set the name 'a' to have value 5
a = 5
```

```python
# Ask if it is equal to the value you think it should be
a == 5
```

---

## [Tutorial] Membership Operators: `in`
The keyword `in` allows you to check for membership: if an element is contained in a larger sequence of elements.

As one use case, `in` can be used to check for **substrings**, i.e., strings that are contained in larger strings:

```python
"hello" in "hello world"
```

```python
"i" in "rhythm"
```

In another use case, `in` can also check whether an array *contains* a certain element:

```python
my_array = make_array(3,2,5,6)
5 in my_array
```

```python
9 in my_array
```

---

## [Tutorial] Compound expressions with Boolean operators `and`, `or`, and `not`

`and`, `or`, and `not` are boolean operators that operate directly on booleans, creating **compound** boolean expressions. For two booleans `x` and `y`:

* `x and y`: Evaluates to `True` when both operands are `True`; otherwise, it evaluates to `False`.
* `x or y`: Evaluates to `False` when both operands are `False`; otherwise, it evaluates to `True`.
* `not x`: Evaluates to the other boolean, i.e., "flips" the value.

| `x` | `y` | `x and y` | `x or y` | `not x` |
| --- | --- | --- | --- | --- |
| `False` | `False` | `False` | `False` | `True` | 
| `False` | `True` | `False` | `True` | `True` |
| `True` | `False` | `False` | `True` | `False` | 
| `True` | `True` | `True` | `True` | `False` |
    
These three operators are useful when writing functions because they help control a function's logic. You will find them useful when writing more complicated functions in the course.

**Tip**: Use parentheses to ensure your boolean operators are being applied correctly and without ambiguity, as below:

```python
(3 * 4 == 12) or (3 + 4 == 50)
```

```python
(3 * 4 == 12) and (3 + 4 == 50)
```

```python
not (3 + 4 == 50)
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Comparisons with NumPy
Learning `np.any()` and `np.all()` is essential since it allows much faster operations in comparison to iterating over the entire array.

Here is the [np.all() documentation](https://numpy.org/doc/stable/reference/generated/numpy.all.html), and [np.any() documentation](https://numpy.org/doc/stable/reference/generated/numpy.any.html).

The function `np.all()` takes in an array and tests whether all array elements evaluate to truthy values. Run the following cells and verify that the output matches what you expect. 

```python
np.all(make_array(True, True, False))
```

```python
np.all(make_array(1, 1, 1, 1))
```

```python
np.all([1, 2.5, -1])
```

Meanwhile, `np.any()` takes in an array and tests whether any array element is evaluated as True. 

```python
np.any(make_array(1, 0))
```

```python
np.any(make_array(False, False, False, False))
```

```python

```

<hr style="border: 1px solid #fdb515;" />

# Part 2: Functions

Writing functions is useful when we want to use the same code repeatedly without wanting to write it out each and every time. When defined, function names by themselves do not execute any code. However, when called, functions operate on arguments and execute lines of code in the function body.

Below is an example of a Python function that doubles its input:

```python
# just run this cell
def double(x):
    return 2 * x
```

Let's break down the different parts of this function:
- The `def` at the beginning tells Python that we are making a new function.
- The blue `double` text is the name of our function. If we want to use this function, we will call it using this name.
- The `x` inside the parentheses is the input to the function, also known as a **parameter**. When we call the function `double`, we pass in different **arguments** to the parameter `x`. A function can have many parameters (even zero), as long as they have different names.
- The `:` indicates we are done defining our function **signature**.
- All lines after this `def` line are the **function body**; these lines are **indented**.
    - You can have as many lines as you want in your functions, but in this case we only have 1.
- The `return` keyword indicates the output to the function, also know as the **return value**. In our case, we want to multiply the parameter `x` by 2 before returning.

```python
# just run this cell
answer = double(5)
answer
```

---

## Question 2.1

Write code that defines the `satisfies` function, which returns `True` if the parameter `x` satisfies *at least one* of the following conditions:
1. Has a value in the range 24 and 32 (inclusive)
2. Is even

Otherwise, if `x` satisfies neither condition, `satisfies` returns `False`. Some examples:

```python
satisfies(27) # True
satisfies(64) # True
satisfies(32) # True
satisfies(33) # False
```

*Hint*: You should try using the modulo operator (%) to help check if a value is odd or even! The `%` operator will provide you with the *remainder* after dividing by a certain number.

Feel free to adjust the example value `val` below to test your function.

```python
def satisfies(x):
    ...

val = 33
satisfies(val)
```

```python
grader.check("q2_1")
```

## Scope

We can also call values that are assigned **outside** of the function, including other pre-existing functions or named values, **inside** a new function that we define. We know this through our current use of Jupyter notebooks: We can run cells sequentially and access values assigned to names in previously-executed cells.

For example, the named `outside_value` can be accessed and read inside the function `greeting`:

```python
outside_value = "Hello, "

def greeting(place):
    return outside_value + place
```

```python
greeting("World!")
```

However, any names assigned **inside** a function cannot be used **outside** the function; this is because these names only exist in the **local scope** of the function, while it is executing.

The locally named `current_year` is only accessible inside the function `age`, and inaccessible otherwise:

```python
def age(birth_year):
    """ This is called a docstring. This explains what the function does:
    Return the age of a person, given their birth year."""
    current_year = 2025
    return current_year - birth_year
```

```python
age(2004)
```

```python
# uncomment the line below, which errors. re-comment if you want to run all cells.
# birth_year
```

---

## Question 2.2

Complete the function `triple_double`, so that it uses the function `double` we used earlier to double an input. Then, take that doubled input and triple it and return the triple-doubled value. 

For example: if we called `triple_double(2)`, the input 2 should be doubled to 4, and then tripled to 12. The function would then output 12.  


```python
def triple_double(x):
    ...
```

```python
grader.check("q2_2")
```

<hr style="border: 1px solid #fdb515;" />

# Part 3: Conditional Statements

When writing functions in Python, we may want the function to behave differently depending on the input---a concept in computer science called **control**. We could choose to write several similar functions to accomplish this, but that would require copying and pasting much of the same code over and over again, only making small changes. Instead, we need a way to tell one singular function to execute different code for different inputs.

A conditional statement is a multi-line statement that allows Python to choose among different alternatives based on the truth value of an expression.

Here is a basic example:

```python
# just run this cell
def sign(x):
    if x > 0:
        return 'Positive'
    else:
        return 'Negative'
```

If the input `x` is greater than `0`, we return the string `'Positive'`. Otherwise, we return `'Negative'`.

```python
sign(3)
```

```python
sign(-2)
```

If we want to test multiple conditions at once, we use the following general format.

```python
if <if expression>:
    <if body>
elif <elif expression 0>:
    <elif body 0>
elif <elif expression 1>:
    <elif body 1>
...
else:
    <else body>
```
        
Only the body for the first conditional expression that is true will be evaluated. Each if and elif expression is evaluated and considered in order, starting at the top. elif can only be used if an if clause precedes it. As soon as a true value is found, the corresponding body is executed, and the rest of the conditional statement is skipped. If none of the if or elif expressions are true, then the else body is executed.

---

## Question 3.1(a)

Water has a different state of matter at different temperatures:

<img src='water.jpeg' width=300>

For a given temperature (in Farenheit), the function `state_of_matter` should return:
* `'ice'`, the solid state of water, below 32°F
* `'water'`, the liquid state of water, between 32°F and 212°F, inclusive
* `'steam'`, the gaseous state of water, above 212°F

Implement the function `state_of_matter` below:

```python
def state_of_matter(temperature):
    if ...:
        ...
    elif ...:
        ...
    else:
        ...
```

```python
grader.check("q3_1_a")
```

We can call the function on different inputs and the function deals with each input differently based on the control logic you gave it. This will be very important when you write functions that deal with large, sometimes unpredictable datasets where the logic of your function will deal with inputs you may not have directly prepared the function for.

Let's use this `state_of_matter` function you defined above to determine the state of water in the following locations during winter:

```python
# just run this cell
city_temps = Table().with_columns(
    "City", make_array("Berkeley", "New York", "Miami", "Earth's Core"),
    "Temperature", make_array(58, 24, 78, 10800)
)
city_temps
```

---

## Question 3.1(b): `apply`

Create a table `city_temps_water`, a new table that copies the two columns of `city_temps` and adds a  third column, `State of Matter`, with the state of water in each location.

_Hint_: Consider using the `apply` method, which returns an array.

```python
city_temps_water = ...
city_temps_water
```

```python
grader.check("q3_1_b")
```

---

## Question 3.3: XOR

The XOR operation on two boolean values is an "exclusive or." The function `xor(x, y)` returns:
* `True` only if exactly one of `x` or `y` is `True` (or truthy).
* `False` otherwise.

Some examples (note that the second-to-last example differentiates an XOR from the `or` operator):

| `x` | `y` | `xor(x, y)` |
| --- | --- | --- |
| `False` | `False` | `False` |
| `False` | `True` | `True` |
| `True` | `False` | `True` | 
| `True` | `True` | `False` |
| `None` | `3` | `True` | 

Implement the `xor` function below.
* As a first pass, try using conditional statements.
* As a challenge, try using only parentheses and boolean operators `and`, `or`, and `not`!

```python
def xor(x, y):
    ...
```

```python
grader.check("q3_3")
```

_Comment_: XOR is a classic paradigm in computer science and hardware design, but implementing it requires a solid understanding of boolean operations. You should feel very proud of yourself for successfully implementing this core function! :-)

# Done! 😇

## Pets of Data 6

You too can be a hero! **Hero** just wanted to let you know!!!

<img src="hero.jpeg" width="50%" alt="photo of dog hero and two other dog friends"/>

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q2_1 (points: 1)
```
>>> isinstance(satisfies(24), bool)
True
```
```
>>> satisfies(27)
True
```
```
>>> satisfies(64)
True
```
```
>>> satisfies(32)
True
```
```
>>> satisfies(33)
False
```


Tests for q2_2 (points: 1)
```
>>> callable(triple_double)
True
```
```
>>> triple_double(2) == 12
True
```
```
>>> triple_double(3) == 18
True
```


Tests for q3_1_a (points: 1)
```
>>> state_of_matter(30)
'ice'
```
```
>>> state_of_matter(95)
'water'
```
```
>>> state_of_matter(100)
'water'
```
```
>>> state_of_matter(220)
'steam'
```


Tests for q3_1_b (points: 1)
```
>>> city_temps_water.num_rows == 4
True
```
```
>>> city_temps_water.num_columns == 3
True
```
```
>>> city_temps_water.sort('Temperature').column('State of Matter').item(0) == 'ice'
True
```
```
>>> city_temps_water.sort('Temperature').column('State of Matter').item(2) == 'liquid'
False
```
```
>>> city_temps_water.sort('Temperature').column('State of Matter').item(-1) == 'steam'
True
```


Tests for q3_3 (points: 1)
```
>>> xor(False, False) == False
True
```
```
>>> xor(False, True)
True
```
```
>>> xor(True, False)
True
```
```
>>> xor(True, True) == False
True
```
```
>>> xor(None, 3) == True
True
```
