```python
# Initialize Otter
import otter
grader = otter.Notebook("lab08.ipynb")
```

<img src="data6.png" style="width: 15%; float: right; padding: 1%; margin-right: 2%;"/>

# Lab 08 – Iteration and Dictionaries

## Data 6, Summer 2026

Welcome to Lab 08! In this lab, we'll be covering **iteration** and **dictionaries**.

```python
# Just run this cell
from datascience import *
import numpy as np
import warnings
warnings.simplefilter('ignore')
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# [Tutorial] Part 1: Iteration with `for` Loops
Iteration is the repeated running of code. Typically, the variables in iteration change sequentially, and the iteration continues to repeat until a boolean condition indicates that the iteration be stopped.

One way to do this is using a `for` loop. A `for` loop begins with the word `for`, followed by a name we want to give each item in the sequence, followed by the word `in`, and ending with an expression that evaluates to a sequence. The indented body of the `for` loop is executed once for each item in that sequence.

For example, say we're given the following array, where each value corresponds to the number of hours of sleep one night of the week: 

```python
hours_of_sleep = make_array(6, 7, 5, 8, 7, 6, 8)
```

We could use a `for` loop to _iterate_ through each value in the array. One simple thing we could do is print each value in the array. 

```python
for num_hours in hours_of_sleep: 
    print(num_hours)
```

Notice that we used `num_hours` as our name for each item in the sequence, but we could have given it another name, and it wouldn't change the functionality of the `for` loop: 

```python
for n in hours_of_sleep: 
    print(n)
```

While the above `for` loop is fairly simple, we could do more than just printing each value. For example: 

```python
night_of_sleep = 1
for num_hours in hours_of_sleep: 
    print("On night " + str(night_of_sleep)+ ", I slept for " + str(num_hours) + " hours.")
    night_of_sleep += 1
```

In the code above, we "initialize" `night_of_sleep` to 1, and use it to "counter" to keep track of what night we are on.

<br></br>
<hr style="border: 1px solid #fdb515;" />

# Question 1 – FizzBuzz

[FizzBuzz](https://en.wikipedia.org/wiki/Fizz_buzz) is a game and common programming challenge. It is described below.

> Players generally sit in a circle. The player designated to go first says the number "1", and the players then count upwards in turn. However, any number divisible by three is replaced by the word "fizz" and any number divisible by five by the word "buzz". Numbers divisible by 15 become "fizz buzz". A player who hesitates or makes a mistake is eliminated from the game.

In this problem, you will write a function that computes the "FizzBuzz sequence." We'll practice using `for` loops to do this.

---
## Question 1.1

First, complete the implementation of the function `single_number`, which takes in an integer (`n`) greater than or equal to 1 and returns the corresponding number or word in the FizzBuzz sequence, **as a string**. For instance, `single_number(2)` should evaluate to `'2'`, and `single_number(30)` should evaluate to `'fizz buzz'`.

_Hint_: Use the remainder operator `%`.<br>
_Hint_: You might find conditionals useful.


```python
def single_number(n):
    ...
```

```python
grader.check("q1_1")
```

---
## Question 1.2

Complete the implementation of the function `fizz_buzz_sequence`, which takes in an integer (`n`) greater than or equal to 1 and returns an array of length `n` containing the first `n` elements of the FizzBuzz sequence.

For instance, `fizz_buzz_sequence(8)` should evaluate to the array `array(['1', '2', 'fizz', '4', 'buzz', 'fizz', '7', '8'])`.

_Hint_: Use `np.append`. For more information on `np.append`, you can reference the `NumPy` documentation [here](https://numpy.org/doc/stable/reference/generated/numpy.append.html). 


```python
def fizz_buzz_sequence(n):
    sequence = make_array()
    for i in np.arange(1, n+1, 1):
        ...
    return sequence
```

```python
grader.check("q1_2")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# [Tutorial] Part 2: Strings

We have already seen strings before in this course, but we want to introduce some new ideas with strings that will be very useful when dealing with data from real-world data sets.

Strings can be viewed as special lists of characters. As we learned in HW04, strings can be printed, sliced, indexed, and more.

One important property about strings is that they are **immutable**, meaning their values cannot be changed once created. For example, say we have the following string: 

```python
our_string = 'flag'
print(our_string)
```

Notice that the following code will result in an error:

```python
our_string[3] = 'x'
```

Instead, we would have to create a new string, using other approaches such as slicing and string concatenation.

```python
our_new_string = our_string[:3] + 'x'
print(our_new_string)
```

Another thing we could do with strings is check if it's _in_ (or _not in_) an array. 

```python
data_classes = make_array('Data 6', 'Data 8', 'Data 100')
"Data 6" in data_classes
```

```python
"Data 8" in data_classes
```

```python
"Data 5" not in data_classes
```

---
## Question 2: Hailstone
The Hofstadter hailstone sequence is a popular number problem. The rules of the sequence are as follows:
1. Pick a positive integer n that is greater than 1 as the start value.
2. If n is even, divide it by 2.
3. If n is odd, multiply it by 3 and add 1.
4. If you continue this process, n will eventually reach 1. 


The integer _n_ will increase and decrease but will always eventually reach 1\*. Analogously, a hailstone travels up and down in the atmosphere before eventually reaching earth.

\*( At least for all numbers that have ever been tried with the sequence. It has yet to be proven that the hailstone sequence will terminate. If this sounds interesting, check out the [halting problem](https://en.wikipedia.org/wiki/Halting_problem) and maybe [CS70](https://www.eecs70.org/).)

_Note_: Skeleton code is provided, you are not required to follow it. 


```python
def hailstone(n): 

    count = ...
    ...
    while ...: 
        ...
        if ...:
            ...
        else: 
            ...
        print(n)
```

```python
grader.check("q2")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Part 3: Comparisons with NumPy
Learning `np.any()` and `np.all()` is essential since it allows much faster operations in comparison to iterating over the entire array.

Here is the [np.all() documentation](https://numpy.org/doc/stable/reference/generated/numpy.all.html), and [np.any() documentation](https://numpy.org/doc/stable/reference/generated/numpy.any.html). 

The function `np.all()` takes in an array and tests whether all array elements evaluate to truthy values. Run the following cells and verify that the output matches what you expect. 

np.all(make_array(True, True, False))

```python
np.all(make_array(1, 1, 1, 1))
```

```python
np.all([1, 2.5, -1])
```

Meanwhile, `np.any()` takes in an array and tests whether any array element is evaluated as True. 

```python
np.any(make_array(1, 0))
```

```python
np.any(make_array(False, False, False, False))
```

```python
# The following line of code errors due to a syntax error. Are you able to spot it?
np.any(make_array'fizz, 'buzz'))
```

# Question 3
How are `np.any()` and `np.all()` equivalent to compound OR and compound AND?


_Type your answer here, replacing this text._

<!-- END QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## [Tutorial] Part 4: Dictionaries


<img src='dictionary.png' width=300>

**Dictionaries** can be very useful. They store key/value pairs that can be used to map one value to another. You can think of a dictionary as a list where the indexes (locations) of the values of the list are no longer their integer locations, but rather their keys.

>In an array, you access the first item with `my_array.item(0)`.

>In a dictionary, you access the "key-th" item with `my_dictionary[key]`.

If we think of list items as having their "address" be their location in the list, then a dictionary value's "address" is its key.

Some important properties of dictionaries to note:
- The key and value **do not** have to be of the same type
- We designate a new key/value entry in a dictionary in this format: *key* **:** *value*
- We store all these key/value entries in a dictionaries with braces `{}` around the ends (like `[]` with a list) and commas separating the entries
- Keys in a dictionary are unique, but values don't have to be unique. 
    - For example {'a': 100, 'a': 200} is not a valid dictionary
    - but, {'a': 100, 'b': 100} is a valid dictionary
    
Let's take a closer look at a dictionary in practice:

```python
my_dictionary = {"a": 100, "b": 200, "c": 300}
print("The value 'a' maps to the value:", my_dictionary["a"])
print("The value 'b' maps to the value:", my_dictionary["b"])
print("The value 'c' maps to the value:", my_dictionary["c"])
```

### How to Access the Data

We can't access a dictionary's values like we can access a list's values. If we want the "first" item in a dictionary, we cannot ask for `my_dictionary[0]`, because this request is really asking "What does the key 0 map to in this dictionary?". If your dictionary does not have a value associated with the key 0, you will get an error.

```python
my_dictionary[0]
```

A `KeyError` warning means that you asked for a key that is not in your dictionary. This may happen when you are writing a function with a dictionary, so if you see it, this is what it means.

### Changing the Data

We can add the key value pair `(key, value)` with the following syntax:

> `my_dictionary[key] = value`

Run the cell below to change the `"d"` entry of the dictionary to 400:

```python
my_dictionary["d"] = 400 # Add the key/value pair ("d", 400) to our dictionary
my_dictionary
```

We can use **any** data type we know as a value in a dictionary...

```python
# Here, the value we add is a list!
my_dictionary["grocery list"] = make_array("apples", "bananas", "carrots")
my_dictionary
```

...including even having a **dictionary itself** as a value! 

```python
my_dictionary["squares"] = {1: 1, 2: 4, 3: 9, 4: 16}
my_dictionary
```

### Dictionary Iteration

We can get a list of a dictionary's keys with the `.keys()` function.

```python
my_keys = my_dictionary.keys()
my_keys
```

```python
# Note the type of this list of keys
type(my_keys)
```

To iterate over the keys in a dictionary, we can use a `for` loop!

```python
for key in my_dictionary:
    print("I am a key, and my name is:", key)
```

We can also get a list of a dictionary's values with the `.values()` function.

```python
my_values = my_dictionary.values()
my_values
```

We can iterate over the values of a dictionary like this:

```python
for value in list(my_dictionary.values()):
    print("I am a value, and my name is:", value)
```

We can use this to do cool things like change all the values in a dictionary!

```python
def add_one_to_dictionary_values(dictionary):
    for key in dictionary:
        dictionary[key] = dictionary[key] + 1
    return dictionary

new_dictionary = {"data": 6, "cs": 61, "poli sci": 1}
modified_dictionary = add_one_to_dictionary_values(new_dictionary)
modified_dictionary
```

We can also use the `.items()` function to return all key-value pairs in a dictionary as a list-like object:

```python
# just run this cell
for k, v in modified_dictionary.items():
    print(k, v)
```

In the cell above, the names `k` and `v` are respectively assigned to each key and value in the `modified_dictionary` dictionary. If you’re curious about how this works, please ask us!

A final note: When provided a dictionary name, Python by default will iterate over the keys of the dictionary:

```python
# just run this cell
for k in modified_dictionary:
	print(k)
```

## Part 4: Dictionary Fundamentals



## Question 4 – Emojify 

The default keyboard on iOS suggests emojis for you to use in place of boring, ordinary words.

<img src = "https://support.apple.com/library/content/dam/edam/applecare/images/en_US/iOS/ios12-iphone-x-messages-replace-words-with-emoji.jpg" width=200>

In this question, you will replicate some of that behavior using dictionaries!


### Emojis in Python
In Python, emojis can be included as part of a string. For example:

```python
'🤗'
```

If you remove the quotes from the emoji above, you will see `SyntaxError: invalid character in identifier`. **Make sure that throughout this question, your emojis are contained within strings!** <br> Fun fact: they cannot currently be used as variable names. Try it and see what error you get!

---
## Question 4a

In the cell below, define a dictionary `fav_emojis` that has the following **five** keys:
- `'happy'`
- `'annoyed'`
- `'tired'`
- `'love'`
- `'food'`

The values corresponding to these five keys must be an emoji. [getemoji.com](https://getemoji.com) allows you to copy and paste emojis. To select an emoji, double click it to highlight it. You may choose any emojis you would like **as long as**:
>- it is copied from the site above
>- it is not in the "New Emojis" category at the bottom

Have fun with it! We've chosen an emoji for `'happy'` for you, but feel free to change it.

**Some troubleshooting tips:**
- After defining your dictionary, you may see some emoijs displayed with `'\U001...'` instead of their actual graphic. When a browser doesn't support the emoji's encoding, the code is displayed instead. If this happens, pick different emojis.
- If you fail the test case that says your emojis are invalid, and you're certain you correctly defined your dictionary, you may consider choosing other emojis that are more generic that are more likely to be recognized by our autograder. This most likely won't be a problem.


```python
fav_emojis = {
    'happy': '😁',
    ...
}

fav_emojis
```

```python
grader.check("q4a")
```

---
## Question 4b

Now, complete the implementation of the function `emojify`, which takes in a string `message` and returns a new string with all instances of any of the keys in `fav_emojis` replaced with their corresponding emoji value. Example behavior is shown below, though the emojis will be different, depending on what you put in `fav_emojis`. If you passed the previous question, you don't need to change your emojis!

```py
>>> emojify('Filing taxes makes me tired and want food.')
'Filing taxes makes me 😵 and want 🌽.'

>>> emojify('I LOVE love life right now. I am so happy – why do you look so annoyed?!')
'I 💋 💋 life right now. I am so 😁 – why do you look so 💀?!'

>>> emojify("It's not you, it's me... I don't make you haPPy, I make you tired.")
"It's not you, it's me... I don't make you 😁, I make you 😵."
```


```python
def emojify(message):
    # This line ensures your code replaces correctly if any of
    # the keys in fav_emojis appears in uppercase in the message
    message = message.lower()

    ...
    
    # Don't change this
    return message
```

```python
grader.check("q4b")
```

### Fun Demo

Run the cell below to produce a text box (don't worry about the code itself). Type text in the text box and watch it get emojified live!

```python
def emojify_live(type_here):
    display(HTML('<h2>' + emojify(type_here) + '</h2>'))
interact(emojify_live, type_here="I LOVE food");
```

<!-- BEGIN QUESTION -->

---
## Question 4c

Nice and simple: What's your favorite emoji? Place it in the Markdown cell below.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

----
## Done! 🤠

## Pets of Data 6

Lila is proud of you! Now that you're done, go to gradescope to turn in this lab.

<img src="lila.jpg" width="50%" alt="A dog with a watermelon helmet."/>

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1_1 (points: 1)
```
>>> callable(single_number)
True
```
```
>>> single_number(30) == 'fizz buzz'
True
```
```
>>> single_number(9) == 'fizz'
True
```
```
>>> single_number(10) == 'buzz'
True
```
```
>>> single_number(1) == '1'
True
```


Tests for q1_2 (points: 1)
```
>>> all(fizz_buzz_sequence(8) == make_array('1', '2', 'fizz', '4', 'buzz', 'fizz', '7', '8'))
True
```
```
>>> all(fizz_buzz_sequence(3) == make_array('1', '2', 'fizz'))
True
```
```
>>> fizz_buzz_sequence(15).item(14) == 'fizz buzz'
True
```
```
>>> sum(fizz_buzz_sequence(1000) == 'fizz') == 267
True
```


Tests for q2 (points: 1)
```
>>> callable(hailstone)
True
```
```
>>> hailstone(2) == 1
2
1
True
```
```
>>> hailstone(3) != 10 and hailstone(10) != 5
3
10
5
16
8
4
2
1
10
5
16
8
4
2
1
True
```
```
>>> hailstone(9) == 19 and hailstone(10) == 6
9
28
14
7
22
11
34
17
52
26
13
40
20
10
5
16
8
4
2
1
10
5
16
8
4
2
1
True
```


Tests for q4a (points: 2)
```
>>> len(fav_emojis) == 5
True
```
```
>>> sorted(list(fav_emojis.keys())) == ['annoyed', 'food', 'happy', 'love', 'tired']
True
```
```
>>> all([len(e) == 1 for e in fav_emojis.values()])
True
```
```
>>> def read_json(path):
...     import json
...     with open(path, 'r') as f:
...         contents = json.load(f)
...     return contents
>>> all_emojis = read_json('emojis.json')
>>> invalid_keys = []
>>> for k in fav_emojis.keys():
...     if fav_emojis[k] not in all_emojis:
...         invalid_keys.append(k)
>>> len(invalid_keys) == 0
True
```


Tests for q4b (points: 4)
```
>>> convert_1 = emojify('Filing taxes makes me tired and want food.').split(' ')
>>> e1_1, e1_2 = (convert_1[4], convert_1[-1][:-1])
>>> e1_1 == fav_emojis['tired'] and e1_2 == fav_emojis['food']
True
```
```
>>> convert_2 = emojify('I LOVE love life right now. I am so happy – why do you look so annoyed?!').split(' ')
>>> e2_1, e2_2, e2_3, e2_4 = (convert_2[1], convert_2[2], convert_2[9], convert_2[-1][:-2])
>>> e2_1 == fav_emojis['love'] and e2_2 == fav_emojis['love'] and (e2_3 == fav_emojis['happy']) and (e2_4 == fav_emojis['annoyed'])
True
```
```
>>> convert_3 = emojify("It's not you, it's me... I don't make you haPPy, I make you tired.").split(' ')
>>> e3_1, e3_2 = (convert_3[9][:-1], convert_3[-1][:-1])
>>> e3_1 == fav_emojis['happy'] and e3_2 == fav_emojis['tired']
True
```
```
>>> convert_4 = emojify('Billy is lovesick, and not the kind that can be solved with food – Billy is annoyed, really.').split(' ')
>>> e4_1, e4_2, e4_3 = (convert_4[2][0], convert_4[12], convert_4[-2][0])
>>> e4_1 == fav_emojis['love'] and e4_2 == fav_emojis['food'] and (e4_3 == fav_emojis['annoyed'])
True
```
