```python
# Initialize Otter
import otter
grader = otter.Notebook("lab10.ipynb")
```

<img src="data6.png" style="width: 15%; float: right; padding: 1%; margin-right: 2%;"/>

# Lab 10 – API Keys

## Data 6, Summer 2026

Welcome to Lab 10! In this lab, we'll be covering **API Keys**.

```python
# Just run this cell
from datascience import *
import numpy as np
import warnings
warnings.simplefilter('ignore')
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Genius API
In this part, you will access the http://genius.com/ lyrics database using the Genius API.

Credit: The Genius API tutorial and exercises are adapted from [Intro to Cultural Analytics](https://melaniewalsh.github.io/Intro-Cultural-Analytics/welcome.html) by [Professor Melanie Walsh](https://melaniewalsh.org/) at the University of Washington.

Note: The below tutorials and Question 1 are identical to the beginning of Project 2. If you have started Project 2, copy your API token from there to complete Question 1.

---

## [Tutorial] Getting a Genius API key 

This tutorial walks you through obtaining an API Key for the Genius API. This tutorial is adapted from [Intro to Cultural Analytics](https://melaniewalsh.github.io/Intro-Cultural-Analytics/welcome.html) by Melanie Walsh and is also available in the [Data 6 Notes](https://data6.org/notes/).

## API Keys

To use the Genius API, you need a special API key, specifically a "Client Access Token", which is kind of like a password. Many APIs require authentication keys to gain access to them. To get your necessary Genius API keys, you need to navigate to the following URL: [https://genius.com/api-clients](https://genius.com/api-clients).

<center><img src="Genius-API.png" width=50% ></center>

You'll be prompted to sign up for [a Genius account](https://genius.com/signup_or_login), which is required to gain API access. Signing up for a Genius account is free and easy. You just need a Genius nickname (which must be one word), an email address, and a password.

Once you're signed in, you should be taken to [https://genius.com/api-clients](https://genius.com/api-clients), where you need to click the button that says "New API Client."

<center><img src="Genius-New-API.png" width=50% ></center>

After clicking "New API Client," you'll be prompted to fill out a short form about the "App" that you need the Genius API for. You only need to fill out "App Name" and "App Website URL."

It doesn't really matter what you type in. You can simply put "Song Lyrics Project" for the "App Name" and the URL for our course website "https://data6.org/su26/" for the "App Website URL."

When you click "Save," you'll be given a series of API Keys: a "Client ID" and a "Client Secret." To generate your "Client Access Token," which is the API key that we'll be using in this notebook, you need to click "Generate Access Token".


---

# Question 1
Include your API key from the Python file `api_key.py`.

Follow the corresponding instructions in the [Data 6 Notes](https://data6.org/notes/18-html/genius.html) to navigate to the `api_key.py` file and edit it.

After you have updated `api_key.py` (the file within this assignment directory), running the below line should load your API token into `client_access_token`.

```python
# before running this cell, make sure that you have updated api_key.py with your API key
%reload_ext autoreload
%autoreload 2

import api_key
# client_access_token = api_key.my_client_access_token
```

```python
grader.check("q1")
```

---
## [Tutorial] Making an API request

## Making an API Request

Making an API request looks a lot like typing a specially-formatted URL. But instead of getting a rendered HTML web page in return, you get some data in return.

There are a few different ways that we can query the Genius API, all of which are discussed in the [Genius API documentation](https://docs.genius.com/#songs-h2). The way we're going to cover in this lesson is [the basic search](https://docs.genius.com/#search-h2), which allows you to get a bunch of Genius data about any artist or songs that you search for:

`http://api.genius.com/search?q={search_term}&access_token={client_access_token}`

Sticking with our Missy Elliott theme/obsession, we're going to search for Genius data about Missy Elliott.

First we're going to assign the string "Missy Elliott" to the variable `search_term`. Then we're going to make an f-string URL that contains the variables `search_term` and `client_access_token`.

```python
search_term = "Missy Elliott"
```

```python
genius_search_url = f"http://api.genius.com/search?q={search_term}&access_token={client_access_token}"
```

This URL is basically all we need to make a Genius API request. Want proof? Run the cell below and print this URL, then copy and paste it into a new tab in your web browser.

```python
genius_search_url
```

It doesn't look pretty, but that's a bunch of Genius data about Missy Elliott!

We can programmatically do the same thing by again using the Python library `requests` with this URL. Instead of getting the `.text` of the response, as we did before, we're going to use `.json()`.

[JSON](https://www.w3schools.com/whatis/whatis_json.asp) is a data format that is commonly used by APIs. JSON data can be nested and contains key/value pairs, much like a Python dictionary.

```python
import requests

response = requests.get(genius_search_url)
json_data = response.json()
```

The JSON data that we get from our Missy Elliott API query looks something like this:

```python
json_data
```

We can index this data (again, like a Python dictionary) and look at the first "hit" about Missy Elliott from Genius.com.

```python
json_data['response']['hits'][0]
```

<br></br>
<hr style="border: 1px solid #fdb515;" />

# [Tutorial] Understanding the JSON from the Genius API call

The Genius API's basic search function returns a JSON, which has data about an artist's **hits**, also known as top songs. There is a lot of information about each hit. This question's goal is to understand the structure of this JSON.


JSON data is much like a **Python dictionary**. We can access values by providing keys. We also note that some of the values are themselves dictionaries, meaning that `json_data` dictionary is a **nested dictionary**. As an example, here are the two entries in the `json_data` dictionary:

* `meta`: Information about the web request, which is itself a dictionary. An HTTP status of 200 means "OK", i.e., data was successfully returned and stored into the JSON response. We generally ignore this key.
* `response`: The data from the web request, which is itself a dictionary. **We mostly care about this `response`.**

As an example, here is the first “hit” about Missy Elliott from Genius.com pertaining to her hit song, "Work It". Before running the cell, check to see that you understand the multiple layers of square brackets. Refer to the full `json_data` structure above.

```
json_data['response']['hits'][0]
              1          2    3
```


1. `json_data['response']`: Within the `json_data` dictionary, look up the key `"response"` and get a dictionary.
2. `json_data['response']['hits']`: Within this dictionary, look up the key `"hits"` and get a list.
3. `json_data['response']['hits'][0]`: Get the first (zero-th) element of the hitlist, which is a dictionary.

```python
# just run this cell
json_data['response']['hits'][0]
```

<br></br>
<hr style="border: 1px solid #fdb515;" />

# Question 2

---
## Question 2a


Genius is a lyrics database. Using the Genius API, we can find the lyrics page URL (i.e., web link) of each song in the provided hits. In the cell below, assign `top_hit_title` to the full title of the first “hit” about Missy Elliott.  A song’s full title is under the key `full_title`. The first hit is the first element in the `hits` list. <br><br>

Use square bracket operations to get this value. Your answer should be in the form:

```
`top_hit_title = json_data[KEY_1][KEY_2][IND][KEY_3][KEY_4]
```

where you provide strings or numbers for `KEY_1`, `KEY_2`, `IND`, `KEY_3`, and `KEY_4`.


```python
top_hit_title = ...
top_hit_title
```

```python
grader.check("q2a")
```

---
## Question 2b

Each hit keeps track of the number of views of its page (i.e., accesses to the hit's URL). In the cell below, assign `top_view_counts` to the page view counts of the first “hit” about Missy Elliott. Hint: Look carefully at the structure of the `json_data` dictionary.

Use square bracket operations to get this value. Your answer should be in the form:

```
top_view_counts = json_data[KEY_1][...]...
```

where you provide a series of square bracket operations, just as in the previous part.

```python
top_view_counts = ...
top_view_counts
```

```python
grader.check("q2b")
```

---
## Question 2c

This question may take a bit longer than the previous parts.

Assign `total_view_counts` to the sum of page view counts across **all** Missy Elliott hits. You may want to loop over all of the hits; we have provided a for loop template below.


```python
total_view_counts = ...
for song in json_data['response']['hits']:
    ...

# do not edit this line
total_view_counts
```

```python
grader.check("q2c")
```

## Pets of Data 6

Bella's feeling pretty stuck this far into the semester. But congrats on finishing this assignment!

<img src="bella.jpg" width="50%" alt="photo of dog stuck in fence"/>

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1 (points: 1)
```
>>> type(client_access_token) == str
True
```


Tests for q2a (points: 1)
```
>>> True == True
True
```


Tests for q2b (points: 1)
```
>>> True == True
True
```


Tests for q2c (points: 3)
```
>>> total_view_counts - total_view_counts <= 1000
True
```
