```python
# Initialize Otter
import otter
grader = otter.Notebook("lab01.ipynb")
```

<img src="data6.png" style="width: 15%; float: right; padding: 1%; margin-right: 2%;"/>

# Lab 1 – Using Jupyter Notebooks

## Data 6, Summer 2026

Welcome to the first lab section of Data 6!

First and foremost, labs are designed to complement homework assignments. We encourage you to work through all questions carefully to familiarize yourself with key programming concepts that you will use later on in your homeworks. Remember, you learn data science by **_doing_** data science!

The goals of this lab section:
* Access all of the technology we'll be using in this class;
* Familiarize yourself with the Jupyter Notebook environment;
* Use arithmetic operators in Python; and
* Work with Python names.

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Part 1: Jupyter Notebook

# Cells

**Not these!** &rarr;

<div align=center>
    
<img src='cell_graphic.jpeg' width=200>
   
</div>

<br>

Jupyter Notebooks are made up of cells. There are two types of cells in a Jupyter Notebook: code cells and [Markdown](https://www.markdownguide.org/getting-started/) cells. 
- Code cells are where we write our Python code.
- Markdown cells allow us to write text, like the text you're reading right now. In Homework 1, you'll get a chance to learn a bit of Markdown.

## Code cells

Running a code cell will execute all of the code it contains.  Any output produced by the code will appear below the cell. Notice the `In:` and `Out:` statements to the left of the cells, followed by a number. That number designates the order in which the cell was run. For example, the number `[3]` next to a cell would indicate that this cell is the third cell that has been run in the notebook so far, or that this is the third time you've run this cell if you've re-run the cell multiple times.

To run the code in a code cell, first click on that cell to activate it.  It'll be highlighted with a little green or blue rectangle. Next, do one of the following:
* **Keyboard shortcut**: Hold down the `Shift` key and press `Return` or `Enter`.
* **Button click**: Click the "Run" button in the cell toolbar above.

While the above two options give no difference in computer runtime performance, it's useful to internalize some common keyboard shortcuts early in your Jupyter experience (think of it like learning "Ctrl-C" or "Command-C" to copy instead of "Right click -> Copy to Clipboard"). While they're not required for this course, keyboard shortuts will greatly improve your quality of life when programming Python (just trust us on this).


We can use code cells to do arithmetic in Python. Try running the code cells below:

```python
3 + 4 ** 2
```

The order of evaluation is what you might expect: parentheses, then exponentiation, then multiplication and division, then addition and subtraction; from left to right. 

```python
2 *(3 + 10)
```

You can use parentheses to modify the typical order of evaluation. **How would the above output change, if you removed the parentheses?**

```python
4 ** 2
1 + 2 + 4 + 8 + 16 + 32 + 64
```

**Note that only the value from evaluating the last expression is displayed;** the output from `4 **2` isn't shown at all. This is expected behavior.

```python
10000
```

Also notice that this cell contains a very simple expression that evaluates to `1000`.

```python
10 + 20 + 30 - max(10, 20, 30) - min(10, 20, 30)
```

The output you see from running these five code cells are the values `19`, `16`, `127`, `1000` and `20` displayed, respectively. 

Comments in code cells are denoted using a hashtag (`#`). Comments are used to explain what code does when necessary (or, occasionally in our case, to provide instructions).

```python
# This is a comment!
19 - 18
```

If you try and use syntax Python doesn't know about, it will give you an error.

```python
3 + * 4
```

```python
# Why does this work while the above cell causes an error?
3 * - 4
```

<br/>

---

## Question 1: Markdown Cells


To edit an existing Markdown cell, double click it and change whatever you want. To *render* (display) it, you can "run" it like you'd run a code cell.

**Edit the following cell to contain your favorite color.**

*Note:* This is how questions will be formatted in your homework assignments.

My favorite color is maroon. *(Type your answer here, replacing the text)*

---

## Question 2: Converting cells

To convert a cell from code to Markdown (or vice versa), you should use the dropdown menu in the cell toolbar. For instance, to convert the following cell from Markdown to code, you should:
1. Click the cell.
2. Click "Markdown" in the cell toolbar, and change it to "Code."

Your cell toolbar should then look like this:

<img src="code-celltoolbar.png" width="50%"/>

**Change the following cell from a Markdown cell to a code cell**, and run it by pressing the "Run" button or using the keyboard shortcut `Shift + Enter`.

3 + 4 + 5  # Change this to a code cell

---

## Adding and deleting cells

To add cells, you can go to "Insert" in the toolbar at the top, then click "Insert cell Above" or "Insert cell Below." To delete cells, you can go to "Edit" in the toolbar at the top and click "Delete cells". (Be very careful when deleting a cell.)

For your convenience, we provide this screenshot that details how to switch between "Edit mode" and "Command mode", which is relevant if you want to learn keyboard shortcuts. Again, keyboard shortcuts are not required for this course, but learning some now will greatly improve your programming quality of life.

<img src='edit-mode.png' width=400>

Some useful keyboard shortcuts:

| Action | Mode | Keyboard shortcut |
| :--- | :--- | :---
| Run cell + jump to next cell | Either (puts you in edit mode) | `SHIFT + ENTER` |
| Save notebook | Either | `CTRL/CMD + S` |
| Switch to command mode | Either (puts you in command mode) | `ESCAPE` |
| Switch to edit mode | Command | `ENTER` | 
| Comment out the current line | Edit | `CTRL/CMD + /` |
| Create new cell above/below | Command |  `A/B` |
| Delete cell | Command | `DD` |
| Convert cell to Markdown | Command | `M` |
| Convert cell to code | Command | `Y` |
| Show all shortcuts | Command |  `H` |


## Question 3: Insert Cell

**Add a Markdown cell below this one, and in it write your middle name, or write "I don't have a middle name".**

```python
# The cell you create should be above this one!
```

## Question 4: Delete Cell

Delete the cell below.

```python
# This should be deleted!
```

Let's take some time to practice applying our knowledge about arithmetic expressions and operators! We have covered several arithmetic operators so far:

```python
(7 - 5) * 2**10 - 3 * 8 + 6 - 8 
```

Let's use these arithmetic expressions to solve the next question!

## Question 5: Checking your code
For this lab, and future labs, you'll start using the built-in tests to check whether your work is correct. Sometimes, there are multiple tests for a single question, and passing all of them is required to receive credit for the question. Please don't change the contents of the test cells.

Go ahead and attempt the question below. Running the cell directly after it will test whether you have assigned `seconds_in_a_decade` correctly. If you haven't, this test will tell you the correct answer. Resist the urge to just copy it, and instead try to adjust your expression. (Sometimes the tests will give hints about what went wrong...)

**Question.** Assign the name `seconds_in_a_decade` to the number of seconds between midnight January 1, 2010 and midnight January 1, 2020. Note that there are two leap years in this span of a decade. A non-leap year has 365 days and a leap year has 366 days.

_Hint_: If you're stuck, the next section shows you how to get hints.


```python
# Change the next line 
# so that it computes the number of seconds in a decade 
# and assigns that number the name, seconds_in_a_decade.

seconds_in_a_decade = ...

# We've put this line in this cell 
# so that it will print the value you've given to seconds_in_a_decade when you run it.  
# You don't need to change this.
seconds_in_a_decade
```

```python
grader.check("q_5")
```


<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Part 2: Integer Division and Names

## Integer Division

Combining several `int`s without using "regular" division always results in another `int`. Below, we use the `type()` function to find the type of the evaluated expression from the previous cell.

```python
type((7 - 5) * 2**10 - 3 * 8 + 6 - 8)
```

As soon as division (`/`) or other `float`s are involved, the result is a `float`:

```python
15 / 3
```

```python
type(15 / 3)
```

We can get `int` results from division by using the **integer division** operator `//`. When the dividend (the number we're dividing) isn't evenly divisible by the divisor (the number we're dividing by), the *result is rounded down*.

```python
15 // 3
```

```python
16 // 3
```

```python
type(15 // 3)
```

We can **cast** `int`s to `float`s and vice versa by applying the `int()` or `float()` functiosn to expressions.

```python
int(15 / 3)
```

```python
float(5)
```

If our `float` is not a round integer, we lose information:

```python
int(4.1)
```

## Names
In natural language, we have terminology that lets us quickly reference very complicated concepts.  We don't say, "That's a large mammal with brown fur and sharp teeth!"  Instead, we just say, "Bear!"

In Python, we do this with *assignment statements*. An assignment statement has a name on the left side of an `=` sign and an expression to be evaluated on the right.

```python
ten = 3 * 2 + 4
```

When you run that cell, Python first computes the value of the expression on the right-hand side, `3 * 2 + 4`, which is the number 10.  Then it assigns that value to the name `ten`.  At that point, the code in the cell is done running.

After you run that cell, the value 10 is bound to the name `ten`:

```python
ten
```

The statement `ten = 3 * 2 + 4` is not asserting that `ten` is already equal to `3 * 2 + 4`, as we might expect by analogy with math notation.  Rather, that line of code changes what `ten` means; it now refers to the value 10, whereas before it meant nothing at all.

If the designers of Python had been ruthlessly pedantic, they might have made us write

    define the name ten to hereafter have the value of 3 * 2 + 4 

instead.  You will probably appreciate the brevity of "`=`"!  But keep in mind that this is the real meaning.

# Question 6: Integers and Names

Su Min, a candy lover, recently returned from a trip to Tokyo, Japan, where she bought a lot of limited edition KitKat flavors. Now, she has 160 KitKats in total! She stores her KitKats in a set of drawers, where each drawer can hold exactly 12 KitKats. Whenever a drawer is full with 12 KitKats, the next KitKat goes into the next drawer. Su Min wants to know how many drawers will be completely filled with 12 KitKats and how many KitKats will be left in the last partially filled drawer.


---

### Question 6a

How many **full** drawers are there? Write an expression that computes the number of full drawers. Your expression shouldn't have any numbers in it; only arithmetic operators and the names `total_kitkats` and `drawer_capacity`. Give the value of your expression the name `full_drawers`. 

_Hint_: Should `full_drawers` be an `int` or a `float`? Why?

```python
total_kitkats = 160
drawer_capacity = 12

full_drawers = ...
full_drawers
```

```python
grader.check("q_6_a")
```


---

### Question 6b

How many leftover KitKats will be in the last **partially full drawer**? Write an expression that computes this value and assign that value to the name `leftover_kitkats`. Your expression shouldn't have any numbers in it; only arithmetic operators and the names `total_kitkats` and `drawer_capacity`.

```python
leftover_kitkats = ...
leftover_kitkats
```

```python
grader.check("q_6_b")
```


---

### Question 6c

How many Kitkats are in **full drawers**? Write an expression that computes this value and assign that value to `kitkats_in_full_drawers`. Your expression shouldn't have any numbers in it; only arithmetic operators and some subset of the names `total_kitkats`, `drawer_capacity`, `full_drawers`, and `leftover_kitkats`.

```python
kitkats_in_full_drawers = ...
kitkats_in_full_drawers
```

```python
grader.check("q_6_c")
```

---

## Comments
You may have noticed these comments in previous questions:

    # Like this
    # Or this
    
This is called a *comment*. It doesn't make anything happen in Python; Python ignores anything on a line after a `#`.  Instead, it's there to communicate something about the code to you, the human reader. Comments are extremely useful. 

<img src="http://imgs.xkcd.com/comics/future_self.png">

Source: [XKCD](http://imgs.xkcd.com/comics/future_self.png)

## Question 8
We’ve been given a line of code, but the variable names aren’t very helpful, which makes it tricky to understand. Your job is to figure out what the code is doing and then rewrite it in a clearer way.

```python
var1 = 212
var2 = (var1 - 32) * 5 / 9 - 273
var2
```

First, think about what kind of real-world calculation this code is making. 

_Hint_: think temperature units: Farenheit, Celsius, and Kelvin.

Then, rewrite the code by breaking it up into step-by-step populations, using descriptive variable names so that it's clearer what we're trying to do.

```python
# Write your code in this cell. 
...
```

<hr>

# Part 3: Policies
Last but not least, please fill out the [Welcome Survey](https://forms.gle/dq74M5MBwLaZJCL99). Please use your Berkeley email to access the form!

**Make `response` = "True" (string) in order to get credit for this lab.**


```python
response = ...
```

```python
grader.check("survey")
```

<hr>

## Done! 😇

Congratulations on finishing your first lab of the semester! Boots and Bear are proud of you
<img src="boots.png" width="50%">

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q_5 (points: [0, 0, 0, 1])
```
>>> 'seconds_in_a_decade' in vars()
True
```
```
>>> seconds_in_a_decade != ...
True
```
```
>>> seconds_in_a_decade != 315360000
True
```
```
>>> seconds_in_a_decade == 315532800
True
```


Tests for q_6_a (points: [0, 1])
```
>>> type(full_drawers) == int
True
```
```
>>> full_drawers == 13
True
```


Tests for q_6_b (points: [0, 1])
```
>>> type(leftover_kitkats) == int
True
```
```
>>> leftover_kitkats == 4
True
```


Tests for q_6_c (points: [0, 1])
```
>>> type(kitkats_in_full_drawers) == int
True
```
```
>>> kitkats_in_full_drawers == 156
True
```


Tests for survey (points: 1)
```
>>> response == True
False
```
