```python
# Initialize Otter
import otter
grader = otter.Notebook("hw03.ipynb")
```

<img src="data6.png" style="width: 15%; float: right; padding: 1%; margin-right: 2%;"/>

# Homework 3 – Tables

## Data 6, Summer 2026

This homework is due on **Tuesday, July 21st at 11:59 PM**. This homework is shorter than other homework assignments to account for the deadline!

You must submit the assignment to [Gradescope](https://www.gradescope.com/). See the [syllabus](https://data6.org/fa25/syllabus) for our homework late submission policy.

**Reminder:** This homework assignment contains hidden tests, which means that even if the test passed, **it does not mean your answer is 100% correct**. The autograder in the homework is intended as a sanity check.

While we encourage you to collaborate with your peers, directly copying solutions is not allowed. Please refer to our [syllabus](https://data6.org/su26/syllabus) on Academic Honesty.
If you find yourself stuck on a problem, we recommend that you make a post on [Ed](https://edstem.org) or attend [office hours](https://data6.org/su26/schedule/).

```python
# just run this cell
from datascience import *
import numpy as np
from IPython.display import YouTubeVideo, HTML, display
from ipywidgets import interact, widgets
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Data 6 Music Dataset

Based on the songs submitted to one of our attendance checks, we have created a dataset containing information about everyone's favorite songs.

```python
# Run this cell!
songs = Table.read_table('songs.csv')
songs
```

Each row is a song submitted through the Data 6 attendance form. We were able to get additional information about each song using the [Spotiy API](https://developer.spotify.com/documentation/web-api), which is a type of software that allows users (us!) to request data. An API stands for _Application Programming Interfaces_. An API allows you to programmatically extract and interact with data from websites like Spotify as well as other social networks, applications, and projects that make their data publicly available, such as Twitter and [The Smithsonian museums](https://www.si.edu/openaccess/devtools).

From the [Data 6 Course Notes](https://data6.org/notes/18-html/genius.html#application-programming-interfaces-apis):

> An API is something that a project or company explicitly designs for data-sharing purposes. Why do companies or projects go to all this trouble? One reason is that it helps to promote the use and further development of an application and its data. For example, Twitter wants other developers to use, integrate, and build upon Twitter tools and data. The Twitter API is the main conduit by which these developers can do so.

Variables:
* `id`: Unique ID for each song, specific to our database (sequentially assigned).
* `track name`: The title of the song.
* `artist name`: The performer/artist of the song. If a song has multiple artists, all names are included, comma-separated.
* `album name`: The name of the album the song is from.
* `genre`: The genre(s) associated with the song's artist(s), pipe-separated (`|`). Note that Spotify's API only exposes genre at the artist level, not the track level, so this reflects the genre tags of the artist(s) rather than the song itself.
* `duration`: The length of the song, in seconds.

## Question 1: Exploratory Data Analysis
Let's look at Data 6 listening habits!


Given the `songs` table, let's investigate the `genre` column. You might have noticed that some songs have multiple genres, and these are separated by the pipe symbol (`|`). For this assignment, we want to look at the `"primary genre"` which is the first genre to appear in the `"genre"` string for each song. 

### Question 1a

Fill in the code below to complete the function `get_primary_genre`. `get_primary_genre` takes in a string of pipe-separated genres and returns the first genre as a `string`. 

For example, `get_primary_genre("latin|bachata")` would return `"latin"`.

```python
def get_primary_genre(genre_string):
    ...
```

```python
grader.check("q1a")
```

#### Question 1b

Let's use our new function! Create a column called "primary genre" containing the primary genre of each song. You can (and should!) use your function `get_primary_genre`. Fill in the code below to add a new column.

```python
songs = ...
```

```python
grader.check("q1b")
```

#### Question 1c

If you look at our songs table, you might notice that there's a lot of columns! In fact, there's so many columns that it's a little difficult to for us to see the new column we just added. 

Part of data "cleaning" means getting rid of columns we don't need. For our exploratory analysis of _genres_ create a new table called `genre_table` which contains the columns `"track name"`, `"genre"`, and `"primary genre"` (in that order!).

```python
genre_table = ...
genre_table
```

```python
grader.check("q1c")
```

Let's explore some of the trends in our data! Run the code below to visualize the most common values for `'primary genre'`.

```python
# Just run this code
genre_table.group('primary genre').barh('primary genre')
```

#### Question 1d

What are the three most popular genres? Assign `top_three` to an array of strings, containing the three most common primary genres in order of popularity descending. Break ties alphabetically in ascending order.

```python
top_three = ...
```

```python
grader.check("q1d")
```

<hr style="border: 1px solid #fdb515;" />

## Part 2: One Hot Encoding

Looking at the bar chart from the previous section, `"primary genre"` contains many categories, several of which have only a handful of songs. `"primary genre"` is stored as a string, which can make it tricky to do arithmatic with (such as summations). We can simplify our genres using [one-hot encoding](https://www.geeksforgeeks.org/machine-learning/ml-one-hot-encoding/) or OHE. One-hot encoding transforms our categorical data into binary variables by creating a new column for each category we care about, where each column contains a `1` if a song belongs to that category ("hot") and a `0` if it doesn't ("cold").

We'll create one-hot encoded columns for "pop", "indie", "r&b", "latin", and "rock". We're focusing on these five genres (rather than all of them) because these sub-genres appear often in `'primary genres'`. For instance, `'r&b'` appears in `'indie r&b'` and `'alternative r&b'`, so encoding `'r&b'` captures a lot of the variation in "primary genre" while keeping the number of new columns manageable. Any song whose primary genre doesn't fall into one of these five will simply have a 0 in every OHE column.

#### Question 2a
It's time to implement OHE! Complete the functions below. Each function takes in a string `genre` and should return `1` if `substring` is contained in `genre`, and `0` otherwise. All six functions follow the same pattern.

_Hint: how can you use typecasting to convert a boolean expression into an integer?_

```python
def is_arabic_hip_hop(genre):
    substring = 'arabic hip hop'
    ...

def is_indie(genre):
    substring = 'indie'
    ...

def is_latin(genre):
    substring = 'latin'
    ...

def is_pop(genre):
    substring = 'pop'
    ...

def is_rock(genre):
    substring = 'rock'
    ...

def is_rb(genre):
    substring = 'r&b'
    ...
```

```python
grader.check("q2a")
```

#### Question 2b

Now that we have our OHE functions, let's use them to add new columns to `genre_table`. Fill in the code below to add the columns `"pop"`, `"indie"`, `"r&b"`, `"latin"`, and `"rock"`. Each column should contain the result of applying the relevant OHE function on the values in `'primary genre'`. Your table should look like the one below:

![reference-table](reference-table.png)

```python

genre_table = genre_table.with_columns(
    'arabic hip hop', ...,
    'indie', ...,
    'latin', ...,
    'pop', ...,
    'rock', ...,
    'r&b', ...,
)
genre_table
```

```python
grader.check("q2b")
```

#### Question 2c

Given our implementation of OHE and the OHE columns — `["pop", "indie", "r&b", "latin", "rock"]` — is it possible for a song to have a `1` in more than one column?

Assign `q2c` to `True` if this is possible, and `False` otherwise.

```python
q2c = ...
```

```python
grader.check("q2c")
```

#### Question 2d

Recall that earlier, we noticed the most popular `'primary genre'` was indie. Now that we've updated our genre classification to include substrings, let's see how many songs are classified as `'indie'` now.

assign `total_indie_songs` to the total number of songs we have encoded as indie.

```python
total_indie_songs = ...
total_indie_songs
```

```python
grader.check("q2d")
```

#### Question 2e

Complete the function `count_songs`, which takes a genre name (as a string) and returns the total number of songs classified under that genre (as an integer). This should be the sum of the corresponding one-hot-encoded column in `genre_table`.

```python
def count_songs(genre):
    ...
```

```python
genres = make_array('pop', 'indie', 'r&b', 'latin', 'rock', 'arabic hip hop')

counts = Table().with_columns('genre', genres)
counts = counts.with_columns('total songs', counts.apply(count_songs, 'genre'))
```

```python
grader.check("q2e")
```

<!-- BEGIN QUESTION -->

#### Question 2f

When we created the column "primary genre", we chose to focus on one piece of data (the primary genre) while setting aside other data (like the secondary genre). Similarly, when we used one-hot encoding to convert genres into binary variables, we selected the most common genres for encoding. As with any decision in data analysis, this choice carries some risk.
    
Take a moment to read [Down With 'Urban'](https://variety.com/2020/music/opinion/music-industry-urban-race-1234628235/), a _Variety_ article discussing the harms of genre classification systems. With this article in mind, what risks come with focusing our analysis on "primary genre"? Include at least one example from the songs dataset to support your answer and your answer should be 3-4 sentences.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

#### Question 2f

Considering your answer to the previous question, (and imagine that we had unlimited time and resources for this homework) what would you recommend doing differently in our music analysis? There is no correct answer here! We'd just like to hear your thought process and insights. Please keep your answer to 2-4 sentences.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<hr style="border: 1px solid #fdb515;" />

## Part 3: Music Recommendations (We have Spotify at home)

Now that we've explored genre classification, let's put it to use! In this section, we'll look at the top 200 songs from 2026 to recommend music to the Data 6 students based on their listening habits.

To get our data, we'll be using [Kaggle](https://www.kaggle.com/), a popular online platform for finding and sharing datasets, hosting data science competitions, and exploring public code notebooks. It's a great resource whenever you need real-world data for a project.

Specifically, we'll be working with the [Spotify Global Top Songs 2026](https://www.kaggle.com/datasets/mkaur1141/spotify-global-top-songs-2026) dataset, which contains the top 200 songs on Spotify globally in 2026, along with information like genre, artist, and other song attributes.

#### Question 3a

Load the file `"music_trends_2026.csv"` in as a table and assign it to `music_trends`. *Hint: look at how we read in `"songs.csv"` earlier in this notebook.*

```python
music_trends = ...
music_trends
```

```python
grader.check("q_3a")
```

#### Question 3b

Complete the function `recommend_top_song`, which takes a `genre` (as a string) and returns a song recommendation from `music_trends`.

- If no songs in `music_trends` match the given genre, return the string `"sorry! this genre isn't in our data set"`.
- Otherwise, find the song that matches the given genre and has the most `'streams'`, and return a string formatted as follows: `"you should listen to {song_title} by {song_artist}!"` (for example: `'you should listen to I Wanna Be Yours by Arctic Monkeys!'`).

We've included skeleton code below — a partial structure with some steps already outlined — to help guide your approach. You're welcome to use it, or write your own function from scratch, as long as it returns the correct output.

```python

def recommend_top_song(genre):
    matches = ...
    if matches.num_rows == 0:
        return "sorry! this genre isn't in our data set"
    genre_songs_sorted = ...
    song_title = ...
    song_artist = ...
    return ...
recommend_top_song('indie rock')
```

```python
grader.check("q3b")
```

#### Question 3c

Real-world music recommendation services (like Spotify) don't just recommend the single most popular song in a genre over and over—they mix in some randomness so listeners get variety, even among songs they already like. Let's build a version of `recommend_top_song` that does this.

To do this, we'll use Python's `random` module, which lets us generate random numbers and make random selections. We're importing it here because we want to randomly select a song from all the matches for a given genre, rather than always returning the top streamed song.

Complete the function `recommend_random_song`, which takes a `genre` (as a string) and returns a **random** song recommendation from `music_trends` that belongs to that `genre`.

- If no songs in `music_trends` match the given genre, return the string `"sorry! this genre isn't in our data set"`.
- Otherwise, use `random.randint` to pick a random row index from the matching songs, and return a string formatted as follows: `"you should listen to {song_title} by {song_artist}!"`. You can read more about Python's `random` library (and `randint`) [here](https://docs.python.org/3/library/random.html).

*Hint: `random.randint(a, b)` returns a random integer between `a` and `b`, inclusive of both endpoints.*

```python

import random
def recommend_random_song(genre):
    matches = ...
    if matches.num_rows == 0:
        return "sorry! this genre isn't in our data set"
    random_index = ...
    random_song = ...
    song_title = ...
    song_artist = ...
    return ...
```

```python
grader.check("q_3c")
```

Try out your function! Re-run the cell below with different values to see what music recommendations are out there.

```python
recommend_random_song('pop')
```

## Pets of Data 6

Portobello and Caramello say great job on finishing this homework! Go submit it to gradescope.

<img src="portobello_caramello.png" width="50%" alt="Two cute guinea pigs"/>

## Submission

Below, you will see two cells. Running the first cell will automatically generate a PDF of all questions that need to be manually graded, and running the second cell will automatically generate a zip with your autograded answers. You are responsible for submitting both the coding portion (the zip) and the written portion (the PDF) to their respective Gradescope portals. **Please save before exporting!**

> **Important: You must correctly assign the pages of your PDF after you submit to the correct gradescope assignment. If your pages are not correctly assigned and/or not in the correct PDF format by the deadline, we reserve the right to award no points for your written work.**

If there are issues with automatically generating the PDF in the first cell, you can try downloading the notebook as a PDF by clicking on `File -> Save and Export Notebook As... -> PDF`. If that doesn't work either, you can manually take screenshots of your answers to the manually graded questions and submit those. Either way, **you are responsible for ensuring your submission follows our requirements, we will NOT be granting regrade requests for submissions that don't follow instructions.**

```python
from otter.export import export_notebook
from os import path
from IPython.display import display, HTML
name = 'hw03'
export_notebook(f"{name}.ipynb", filtering=True, pagebreaks=True, exporter_type = "html")
if(path.exists(f'{name}.pdf')):
    display(HTML(f"Download your PDF <a href='{name}.pdf' download>here</a>."))
else:
    print("\n Pdf generation failed, please try the other methods described above")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1a (points: [0])
```
>>> type(get_primary_genre('latin|bachata')) == str
True
```


Tests for q1b (points: [0])
```
>>> type(get_primary_genre('latin|bachata')) == str
True
```


Tests for q1c (points: [0.5, 0.5])
```
>>> genre_table.num_columns == 3
True
```
```
>>> genre_table.labels == ('track name', 'genre', 'primary genre')
True
```


Tests for q1d (points: [0])
```
>>> type(top_three) == np.ndarray
True
```


Tests for q2a (points: [1])
```
>>> all([is_arabic_hip_hop('arabic hip hop') == 1, is_indie('indie') == 1, is_latin('latin') == 1, is_pop('pop') == 1, is_rock('rock') == 1, is_rb('r&b') == 1])
True
```


Tests for q2b (points: [0, 1])
```
>>> all([label in genre_table.labels for label in ['arabic hip hop', 'indie', 'latin', 'pop', 'rock', 'r&b']])
True
```
```
>>> all([genre_table.num_rows == len(genre_table.column(label)) for label in ['arabic hip hop', 'indie', 'latin', 'pop', 'rock', 'r&b']])
True
```


Tests for q2c (points: [0])
```
>>> type(q2c) is bool
True
```


Tests for q2d (points: [0])
```
>>> type(total_indie_songs) is int
True
```


Tests for q2e (points: [0.5, 0.5])
```
>>> count_songs('r&b') == 5
True
```
```
>>> count_songs('indie') == 6
True
```


Tests for q3b (points: [1, 1])
```
>>> recommend_top_song('not a valid genre') == "sorry! this genre isn't in our data set"
True
```
```
>>> recommend_top_song('indie rock') == 'you should listen to I Wanna Be Yours by Arctic Monkeys!'
True
```


Tests for q_3a (points: [0])
```
>>> isinstance(music_trends, Table)
True
```


Tests for q_3c (points: [1])
```
>>> recommend_top_song('not a valid genre') == "sorry! this genre isn't in our data set"
True
```
