```python
# Initialize Otter
import otter
grader = otter.Notebook("hw05.ipynb")
```

<img src="data6.png" style="width: 15%; float: right; padding: 1%; margin-right: 2%;"/>

# Homework 5 — Functions, Iteration, and Measurements

## Data 6, Summer 2026

In this homework assignment, you will apply what you've learned about functions, conditionals and iteration to write more advanced algorithms. 

This homework is due on **Monday, August 3rd at 11:59PM**. There is a 24-hour grace period for this assignment.

You must submit the assignment to [Gradescope](https://www.gradescope.com/). See the [syllabus](https://data6.org/su26/syllabus) for our homework late submission policy.

**Reminder:** This homework assignment contains hidden tests, which means that even if the test passed, **it does not mean your answer is 100% correct**. The autograder in the homework is intended as a sanity check.

While we encourage you to collaborate with your peers, directly copying solutions is not allowed. Please refer to our [syllabus](https://data6.org/su26/syllabus) on Academic Honesty.
If you find yourself stuck on a problem, we recommend that you make a post on [Ed](https://edstem.org) or attend [office hours](https://data6.org/su26/schedule/).

```python
# just run this cell
from datascience import *
import numpy as np
from IPython.display import YouTubeVideo, HTML, display
from ipywidgets import interact, widgets
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />


# Part 1: Iteration with Loops

Recall that the `join` string method concatenates the elements in a list (or array) of strings, separated by the provided string itself:

```python
arr = make_array("Go", "Bears", "!")
'...'.join(arr)
```

<!-- BEGIN QUESTION -->

The process of joining strings is one example of the **fencepost problem** we saw in class: To join N strings together, we need to concatenate with N-1 separators (here, the ellipsis `'...'`).


<hr style="border: 1px solid #fdb515;" />

# Question 1: Fenceposts and Strings

Implement the function `join_strings` below, which joins an array of strings `string_arr` with the separator string `sep`.

Example: `join_strings("...", make_array("Go", "Bears", "!")` should return `'Go...Bears...!'`.

**IMPORTANT**: **You should not use the `join` string method.** Instead, your implementation should use a **for loop.** 

_Hints_:
* For this fencepost problem, the "fenceposts" or "boundaries" are the elements of `string_arr`, and the "fences" are the `sep` strings. There is one less fence than fencepost.
* For an N-element array, how do we iterate over the first N-1 elements? `np.arange` and `len` might be useful!
* To concatenate strings, you can use the `+` operator.
* Optional shorthand: `result_str += "test"` is equivalent to `result_str = result_str + "test"`, provided that `result_str` is a string.

```python
def join_strings(sep, string_arr):
    result_str = ""

    result_str += ... # you may not need this
    for ... in ...:
        result_str += ...
        result_str += ...
    
    result_str += ... # you may not need this
    return result_str

join_strings("...", make_array("Go", "Bears", "!"))
```

```python
grader.check("q1")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->


<hr style="border: 1px solid #fdb515;" />

# Question 2: While loops

Implement the function `less_than_or_equal_to` that takes in a number `val` and a **sorted array** of numbers `number_arr` and returns a **new sorted array** containing only the numbers **less than or equal to** `val`. Note that if `val` is less than all of the numbers in `number_arr`, then return an empty array; if it is greater than all numbers, then return a copy of `number_arr`.

Example: `less_than_or_equal_to(3, make_array(1, 2, 3, 3, 4))` should return `array([ 1.,  2.,  3.,  3.])`.

**IMPORTANT**: Your implementation should use a **while loop**, not a for loop.

_Hints_:
* You may find it useful to iteratively build the return array. You may want to use `np.append`; see the [Data 6 Python Reference](https://data6.org/notes/reference) and the lecture notes for ideas.
* Because you can assume that `number_arr` is sorted, you can stop checking elements once they are larger than `val`.
* Unlike for loops, while loops do not automatically assign a name that iterates across an array (an iteration variable). Instead, you will need to keep track of values outside of the loop. We've provided one suggested approach below, but many are possible. Feel free to edit/remove/add as needed.
* Our solution uses the Boolean operator `and` in our while condition _and_ leverages the Boolean shortcircuiting strategy discussed in previous assignments. Depending on your implementation, you may have something different.

```python
def less_than_or_equal_to(val, number_arr):
    return_arr = ...
    i = ...             # feel free to edit/remove/add
    
    while ...:
        ...
        ...
    return return_arr

less_than_or_equal_to(3, make_array(1, 2, 3, 3, 4))
```

```python
grader.check("q2")
```

<!-- END QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />


# Part 2: Billboard Charts 📈


Music artists and fans alike like to keep track of the **longest #1 streak**, i.e., the most consecutive weeks a song has been ranked #1 on music charts. This statistic is trickier to compute than we may think. This question suggests one possible approach. Try it out!

<br></br>
<hr style="border: 1px solid #fdb515;" />

# Question 3: Exploratory Data Analysis
Run the cell below to load in data from the Billboard Top 100 music charts in 2025. If you're unfamiliar with the *Billboard Top 100*, you can read about the chart [here](https://www.billboard.com/charts/).

```python
# Run this cell!
billboard = Table.read_table('billboard-2025.csv')
billboard
```

Each row is a song's weekly ranking on the Billboard Top 100 chart. Variables:
* `chart_week`: The chart week date; this string (YYYY-MM-DD) is the last day of that corresponding week.
* `current_week`: The current rank/position of this song on the Billboard Top 100 chart.
* `title`: The title of the song.
* `performer`: The performer/artist of the song.
* `last_week`: The rank/position of this song on last week's charts.
* `peak_pos`: The peak or highest rank/position this song has reached, up to and including this week.
* `wks_on_chart`: The number of weeks this song has ever been in the Billboard Top 100 chart, up to and including this week. For example, Ariana Grande's "Santa Tell Me" (released in late 2014) has been on the Billboard Top 100 chart 19 times total but when we consider 2025 data only charts once.

In this question, we'll focus on two songs from 2025 that have charted #1 at some point:
* "Golden" by HUNTR/X (EJAE, Audrey Nuna & Rei Ami). [YouTube link](https://www.youtube.com/watch?v=yebNIHKAC4A)
* "Luther" by Kendrick Lamar and SZA. [YouTube link](https://www.youtube.com/watch?v=sNY_2TEmzho)

```python
# just run this cell and the next one
golden = billboard.where('title', 'Golden').sort('chart_week')
golden.show(18)
```

```python
# just run this cell and the previous one
luther = billboard.where('title', 'Luther').sort('chart_week')
luther
```

<!-- BEGIN QUESTION -->

---

## Question 3a

Consider the `golden` table, which includes weeks that the song "Golden" by HUNTR/X charted on the Billboard Top 100. Using Table methods, find the number of weeks that the song "Golden" by HUNTR/X ranked #1 in 2025. Assign this value to `golden_one_weeks`.

```python
golden_one_weeks = ...
golden_one_weeks
```

```python
grader.check("q3a")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

---

## Question 3b

Is the `golden_one_weeks` value **greater than**, **less than**, or **equal to** the "longest #1 streak" statistic we are looking to compute?

Assign the value `q3b` to one of the strings `"greater"`, `"less"`, or `"equal"` corresponding to your answer.

```python
q3b = ...
```

```python
grader.check("q3b")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

---

## Question 3c

Explain your answer to the previous part in no more than 1-2 sentences. Your answer must start with "Higher.", "Lower.", or "Equal.", then follow with justification.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<hr style="border: 1px solid #fdb515;" />

# Question 4: Longest #1 Streak

To compute the "longest #1 streak" statistic, we must loop through each week of a song and keep track of the longest streak seen so far. This "tracking" is relatively tricky, so we will first solve a related problem, then tackle our original statistic.

First, let's create two new arrays corresponding to the *weekly ranks* (in chronological order) of "Golden" and "Luther", `golden_ranks` and `luther_ranks`:

```python
# just run this cell and the next one
golden_ranks = billboard.where('title', 'Golden').sort('chart_week').column("current_week")
golden_ranks
```

```python
# just run this cell and the previous one
luther_ranks = billboard.where('title', 'Luther').sort('chart_week').column("current_week")
luther_ranks
```

<!-- BEGIN QUESTION -->

---

## Question 4a

Let's first solve a related problem. Implement the function `get_lowest_pos`, which takes in an array of weekly, chronologically ordered ranks for a particular song (e.g., `golden_ranks`) and returns the _lowest_ rank on the Billboard Top 100 that this song achieved. For example, "Golden"'s lowest rank (out of 100) was 81, and "Luther"'s lowest rank (out of 100) was 38.

**IMPORTANT**: Do **not** use `max` or `min` or any array methods. Do **not** use Table methods on `billboard`. The point of this question is to compute this value with a for-loop by iterating over the values of `rank`. This will help you practice keeping track of an intermediate value in order to solve a more complicated problem.

_Hints_:
* To help you, we have created an intermediate value named `lowest_pos` declared outside of the for loop; this tracks the current lowest ranking seen so far.
* In each iteration, you should update `lowest_pos` if the current week's ranking is _below_ the current lowest ranking seen so far (are lower ranks smaller or larger values?). This new value of `lowest_pos` is then used for the comparison for later iterations—until it too gets replaced.
* In each iteration, if the current week is not below the current lowest ranking, you should not update anything.

```python
def get_lowest_pos(ranks):
    lowest_pos = ranks.item(0)
    for rank in ranks:
        ...
    ...

# do not change the lines below
print('Lowest "Golden" rank: ', get_lowest_pos(golden_ranks))
print('Lowest "Luther" rank: ', get_lowest_pos(luther_ranks))
```

```python
grader.check("q4a")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->


---
## Question 4b

Below, complete the implementation of the function `get_longest_one_streak`, which takes in an array of weekly, chronologically ordered ranks for a particular song (e.g., `golden_ranks`) returns the longer number of weeks that song was consecutively ranked first. Examples:
* `get_longest_one_streak(golden_ranks)` returns 7.
* `get_longest_one_streak(luther_ranks)` returns 23.
* `get_longest_one_streak(make_array(17, 2, 2, 2, 3, 2, 2, 1, 4, 3, 1, 1, 1, 1))` returns 4.

_Hints_:
* Consider the previous example, where we kept track of an intermediate value and reset it when the current element warranted it. Here, we keep track of *two* intermediate values:
    * The current streak of #1 ranks (assuming we are in one), and
    * The longest streak of #1 ranks we have seen thus far.
* If the current week's rank is 1, we are in a "#1 streak", so we should increment our current streak's count.
* If the current week's rank is not 1, we are not in a #1 streak. If we just ended our current #1 streak, we should update the longest #1 streak if needed, and reset our current #1 streak back to zero.
* Depending on your solution, this may also be an instance of a **fencepost problem**. Consider the last test case above, where the longest streak occurs at the end of the array.

```python
def get_longest_one_streak(ranks):
    longest_one_streak = ...
    current_one_streak = ...
    for num in ranks:
        ...

    return ...

# do not change the lines below
print("Longest #1 Streak")
print('"Golden": ', get_longest_one_streak(golden_ranks))
print('"Luther": ', get_longest_one_streak(luther_ranks))
```

```python
grader.check("q4b")
```

<!-- END QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />


# Part 3: Body Mass Index

From the [National Institute of Health (NIH)](https://www.nhlbi.nih.gov/calculate-your-bmi): 
> **Body mass index (BMI)** is a measure of body fat based on height and weight that applies to adult men and women. Your BMI is just one piece of the puzzle. It’s based on height and weight but doesn’t take into account your muscle mass, bone density, or body composition. Your healthcare provider will consider whether your BMI is too high or too low for you.

The last quoted sentence packs a lot of information and is the focus of this section, where we analyze **BMI as a measure of health**.
**BMI** is defined according to the following formula (where 703 comes from converting to the correct units, kilograms per square meter):

$$ \text{BMI} \approx \frac{\text{mass in lb}}{(\text{height in inches})^2} \times 703 $$

and **BMI categories** (as per the [World Health Organization](https://www.who.int/data/gho/data/themes/topics/topic-details/GHO/body-mass-index)) are different body fat classes based on BMI ranges. For adults over 20:

| BMI Range        | BMI Category        |
|------------------|---------------------|
| Less than 18.5   | Underweight         |
| 18.5 to <25      | Healthy             |
| 25 to <30        | Overweight          |
| 30 or greater    | Obesity             |

These BMI categories are used widely in United States healthcare systems to predict a patient's obesity. Obesity is associated with health risks such as diabetes; therefore, it is useful to have a predictor for a person's obesity for subsequent treatment and eligibility for different medical procedures. However, we shall see that BMI (and its categories) falls short when it comes to **measurement reliability and validity**.

Let's consider a dataset of what (American) society would classify as relatively healthy individuals: professional athletes. Load in the [ESPN](https://www.espn.com/nfl/team/roster/_/name/sf/sort/experience/dir/desc/san-francisco-49ers) dataset of the 2025 roster of the 49ers, San Francisco's professional football team:

```python
# just run this cell
football_roster = Table.read_table("roster_49ers.csv")
football_roster.show(5)
```

---

We can use the `HT` and `WT` columns to compute BMI, but first we must pre-process these strings into integer values.

## [Tutorial] `s.strip()` and more about `join`

To do so, we will introduce one more useful string method: `strip`, which trims off characters off the left and right side of a given string. `s.strip()` returns a new string.

Run the cells below to see how `strip` works. Recall that the escape character `\` can help represent single and double quotes, e.g., `"\""` is a length-one string with one double quote.

```python
height_str = "6' 3\""
print(height_str)
```

```python
print(height_str.strip("\""))
```

Without any arguments, `strip` trims off whitespace:

```python
height_str_space = "     6' 3\"        "
print("string with whitespace:")
print("<begin>" + height_str_space + "<end>")
```

```python
print("<begin>" + height_str_space.strip() + "<end>")
```

The methods `join` and `strip` can help us extract the feet and inches from a height string. Using square-bracket indexing (though the `item` array method would have worked fine, too):

```python
# just run this cell
height_str_split = height_str.split("\' ")
print("array:", height_str_split)

foot_str = height_str_split[0]
inches_str = height_str_split[1].strip("\"")
print("feet:", foot_str)
print("inches:", inches_str)
```

<!-- BEGIN QUESTION -->

---

## Question 5a

Complete the `height_in_inches` function below, which takes in a height string, e.g., `6' 3"`, and returns the height as a floating point value **in inches**, e.g., `75.0`.

The tutorial above may be useful, though you may need to cast strings to floats with `float()`. You may need multiple lines; our approach uses two.

```python
def height_in_inches(height_str):
    ...
    return ...

height_in_inches("6' 3\"")
```

```python
grader.check("q5a")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

---

## Question 5b

Complete the `weight_in_pounds` function below, which takes in a weight string, e.g., `220 lbs`, and returns the weight as a **floating point** value **in pounds**, e.g., `220.0`.

_Hint_: Use `strip` and cast to `float`. Remember to `return` a value.

```python
def weight_in_pounds(weight_str):
    ...
```

```python
grader.check("q5b")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

---

## Question 5c

Complete the `compute_bmi` function below, which takes in a height string, e.g., `6' 3"` and a weight string, e.g., `220 lbs`, and returns the BMI computed as a float value, e.g., `27.49511111111111`. Do not round.

**BMI** is defined according to the following formula (where 703 comes from converting to the correct units, kilograms per square meter):

$$ \text{BMI} \approx \frac{\text{mass in lb}}{(\text{height in inches})^2} \times 703 $$

_Hint_" You may find useful the two functions you defined earlier: `height_in_inches` and `weight_in_pounds`.

```python
def compute_bmi(height_str, weight_str):
    height_in = ...
    weight_lb = ...
    ...

compute_bmi("6' 3\"", "220 lbs")
```

```python
grader.check("q5c")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

---

## Question 5d

Finally, create two items:
1. The `classify_bmi` function, which takes a BMI value and returns the BMI category as a string e.g., `"Overweight"` if the argument is less than 18.5, or `"Healthy"` if the argument is between 18.5 (inclusive) and 25 (exclusive), and so on.
1. The `football_roster_bmi` table, which is a copy of the original `football_roster` table with two new columns: `BMI` (the BMI of each individual) and `BMI Category` (the BMI classification based on the individual's `BMI`).

From before, **BMI categories** (as per the [World Health Organization](https://www.who.int/data/gho/data/themes/topics/topic-details/GHO/body-mass-index)) are different body fat classes based on BMI ranges. For adults over 20:

| BMI Range        | BMI Category        |
|------------------|---------------------|
| Less than 18.5   | Underweight         |
| 18.5 to <25      | Healthy             |
| 25 to <30        | Overweight          |
| 30 or greater    | Obesity             |

The first few rows of `football_roster_bmi`:

| Team | Name | Number | POS | Age | HT	| WT	| Exp	| College	| BMI	| BMI Category |
| --- | --- | --- | --- | --- | --- | --- | --- | --- | --- | --- |
| Offense	| Mac Jones	| 10	| QB	| 27	| 6' 3"	| 220 lbs	| 5		| Alabama	| 27.4951	| Overweight |
| Offense	| Adrian Martinez	| 4	| QB	| 25	| 6' 2"	| 220 lbs	| 1	| Kansas State	| 28.2432	| Overweight |


_Hints_:
* For the `BMI` column, you may find useful the function you defined earlier: `compute_bmi`.
* For the `BMI Category` column, you may find useful the `classify_bmi` function that you will implement.
* You will need the `apply` Table method in conjunction with `with_column` to build the two new columns.

```python
def classify_bmi(bmi):
    # TODO: implement this function!
    if ...:
        ...
    elif ...:
        ...
    elif ...:
        ...
    ...

# TODO: make the new table!
football_roster_bmi = ...
football_roster_bmi = ...
football_roster_bmi.show(2)
```

```python
grader.check("q5d")
```

<!-- END QUESTION -->


<hr style="border: 1px solid #fdb515;" />

# Question 6: Reliability and Validity of BMI as a Measure of Health

Consider the implications of the BMI classification you performed above. Run the cell to count how many players are in each BMI Category.

```python
# just run this cell
football_roster_bmi.group("BMI Category")
```

<!-- BEGIN QUESTION -->

---

## Question 6a

In the cell below, evaluate if **BMI is a _valid_ measure of health**. Answer yes/no, then justify your answer in 3-5 sentences based on your understanding of measurement reliability/validity, the NIH description of BMI, and your analysis above.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

Next, let's consider if BMI is a **_reliable_** measure of health.

Run the code cell below. It will present you with BMI classifications of provided heights and weights, which you can toggle with the sliders. Pay special attention to "**boundary conditions**," where an individual's BMI classification can fluctuate between BMI categories, e.g., "Healthy" and "Overweight" due to small changes in height or weight.

If the below cell produces an error, verify that your prior function implementations are working:
* `compute_bmi` should be able to take as arguments non-integer heights and weights. This may require debugging `height_in_inches` and `weight_in_lbs`.
* `classify_bmi` should be able to take as arguments non-integer BMI values.

```python
# just run this cell and play with the sliders,
# paying special attention to boundary conditions
from ipywidgets import interactive

def interactive_bmi(feet, inches, lbs):
    height_str = f"{feet}' {inches}\""
    weight_str = f"{lbs} lbs"
    bmi = compute_bmi(height_str, weight_str)
    bmi_category = classify_bmi(bmi)
    print("Height:      ", height_str)
    print("weight:      ", weight_str)
    print("BMI:         ", bmi)
    print("BMI Category:", bmi_category)

interactive(interactive_bmi, feet=(5, 6, 1), inches=(0, 12, 0.1), lbs=(100, 300, .5))
```

<!-- BEGIN QUESTION -->

---

## Question 6b

On average, an individual's weight and height measurements can respectively fluctuate 2-5 lbs and 0.5-1 inch, on any given day.

In the cell below, evaluate if **BMI is a _reliable_ measure of health**. Answer yes/no, then justify your answer in 3-5 sentences based on your understanding of measurement reliability/validity, the NIH description of BMI, and your analysis above. Further consider how height and weight data may be collected—not just for professional athletes, but regular individuals looking to assess their BMI for any number of reasons.

<!-- END QUESTION -->


<hr style="border: 1px solid #fdb515;" />

# Question 7 - BMI Critiques and Alternatives

The Body Mass Index (BMI) is widely integrated into the healthcare systems of the United States. However, its use is publicly debated: critics argue that the measure oversimplifies complex health and social factors; when used as a screening tool, it is an ineffective proxy for making healthcare decisions for individuals.

Read the following articles on BMI provides a particular historical and current context:

* ["Why BMI is a flawed health standard, especially for people of color."](https://www.washingtonpost.com/lifestyle/wellness/healthy-bmi-obesity-race-/2021/05/04/655390f0-ad0d-11eb-acd3-24b44a57093a_story.html) *Washington Post*, May 2021. [Drive Link](https://drive.google.com/file/d/1W9mKym-0r1AXtVUlu4NT6i7fKW0abbW9/view?usp=drive_link)
* [“Why Body Mass Index (BMI) Can Be Problematic for the BIPOC Community”](https://www.everydayhealth.com/weight/why-body-mass-index-bmi-can-be-problematic-for-the-bipoc-community/). Everyday Health, January 2025. 

---
   
## Question 7a

According to the articles above, which of the below statements are correct takeaways? For each of the parts below, assign the provided name in the cell to the corresponding boolean value `True` or `False`.

<!-- BEGIN QUESTION -->

### Question 7a(i)

BMI as a measure was originally established to measure the height and weight of an "average" White, European man.

```python
article_i = ...
```

```python
grader.check("q_7ai")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->


### Question 7a(ii)

In the U.S. healthcare system, a patient's BMI is computed and checked against thresholds called "cut points," which are then used to determine eligibility for particular medications, surgeries, and even certain types of military service.

```python
article_ii = ...
```

```python
grader.check("q_7aii")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->


### Question 7a(iii)

Health institutions in the United States have universally adopted different cut points based on race and ethnicity, including for Latino people, Black people, and people of Asian descent.

```python
article_iii = ...
```

```python
grader.check("q_7aiii")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 7b

According to the articles above, what are potential alternatives to BMI when it comes to assessing an individual’s risk of obesity? Select all that that apply:

A. Socioeconomic status

B. Waist circumference

C. Race/ethnicity

D. Smoking status

E. Place of residence

F. Age

G. Sex

Assigning `bmi_alternatives` to an array of valid letter choices. For example, if you want to pick choices A, and G, `bmi_alternatives = make_array("A", "G")`.

```python
bmi_alternatives = ...
bmi_alternatives
```

```python
grader.check("q_7b")
```

<!-- END QUESTION -->

# Done!

Congrats!

The point breakdown for this assignment is given in the table below:

| **Category** | Points |
| --- | --- |
| Autograder | 29 |
| Written  | 5 |
| **Total** | 34 |

## Pets of Data 6

Great job on the homework! Now that you've finished the homework, Ranger has decided to take a nap. While he's sleeping go submit this notebook to Gradescope!

<img src="ranger.png" width="50%" alt="A black and white dog sleeping with his tongue out."/>

## Submission

Below, you will see two cells. Running the first cell will automatically generate a PDF of all questions that need to be manually graded, and running the second cell will automatically generate a zip with your autograded answers. You are responsible for submitting both the coding portion (the zip) and the written portion (the PDF) to their respective Gradescope portals. **Please save before exporting!**

> **Important: You must correctly assign the pages of your PDF after you submit to the correct gradescope assignment. If your pages are not correctly assigned and/or not in the correct PDF format by the deadline, we reserve the right to award no points for your written work.**

If there are issues with automatically generating the PDF in the first cell, you can try downloading the notebook as a PDF by clicking on `File -> Save and Export Notebook As... -> PDF`. If that doesn't work either, you can manually take screenshots of your answers to the manually graded questions and submit those. Either way, **you are responsible for ensuring your submission follows our requirements, we will NOT be granting regrade requests for submissions that don't follow instructions.**

```python
from otter.export import export_notebook
from os import path
from IPython.display import display, HTML
name = 'hw05'
export_notebook(f"{name}.ipynb", filtering=True, pagebreaks=True, exporter_type = "html")
if(path.exists(f'{name}.pdf')):
    display(HTML(f"Download your PDF <a href='{name}.pdf' download>here</a>."))
else:
    print("\n Pdf generation failed, please try the other methods described above")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1 (points: 3)
```
>>> callable(join_strings)
True
```
```
>>> type(join_strings('...', make_array('Go', 'Bears', '!'))) == str
True
```
```
>>> len(join_strings('...', make_array('Go', 'Bears', '!')).split('...')) == 3
True
```
```
>>> join_strings('...', make_array('Go', 'Bears', '!')) == 'Go...Bears...!'
True
```


Tests for q2 (points: 3)
```
>>> callable(less_than_or_equal_to)
True
```
```
>>> len(less_than_or_equal_to(3, make_array(1, 2, 3, 3, 4))) > 0
True
```
```
>>> 1 in less_than_or_equal_to(3, make_array(1, 2, 3, 3, 4))
True
```
```
>>> 3 in less_than_or_equal_to(3, make_array(1, 2, 3, 3, 4))
True
```
```
>>> np.all(less_than_or_equal_to(3, make_array(1, 2, 3, 3, 4)) <= 3)
True
```
```
>>> len(less_than_or_equal_to(0, make_array(1, 2, 3, 3, 4))) == 0
True
```
```
>>> secret_arr = make_array(1, 2, 3, 3, 4)
>>> len(less_than_or_equal_to(5, secret_arr)) == len(secret_arr)
True
```


Tests for q3a (points: 1)
```
>>> isinstance(golden_one_weeks, int)
True
```
```
>>> golden_one_weeks == 8
True
```


Tests for q3b (points: 1)
```
>>> isinstance(q3b, str)
True
```
```
>>> q3b in ['greater', 'less', 'equal']
True
```


Tests for q4a (points: 1)
```
>>> callable(get_lowest_pos)
True
```
```
>>> secret_golden_ranks = billboard.where('title', 'Golden').sort('chart_week').column('current_week')
>>> get_lowest_pos(secret_golden_ranks) == 81
True
```
```
>>> secret_luther_ranks = billboard.where('title', 'Luther').sort('chart_week').column('current_week')
>>> get_lowest_pos(secret_luther_ranks) == 38
True
```


Tests for q4b (points: 5)
```
>>> callable(get_longest_one_streak)
True
```
```
>>> secret_golden_ranks = billboard.where('title', 'Golden').sort('chart_week').column('current_week')
>>> get_longest_one_streak(secret_golden_ranks) == 7
True
```
```
>>> secret_luther_ranks = billboard.where('title', 'Luther').sort('chart_week').column('current_week')
>>> get_longest_one_streak(secret_luther_ranks) == 13
True
```


Tests for q5a (points: 2)
```
>>> callable(height_in_inches)
True
```
```
>>> isinstance(height_in_inches('6\' 3"'), float)
True
```
```
>>> height_in_inches('6\' 3"') == 75.0
True
```
```
>>> height_in_inches('7\' 2"') == 86.0
True
```
```
>>> height_in_inches('7\' 2.5"') == 86.5
True
```


Tests for q5b (points: 2)
```
>>> callable(weight_in_pounds)
True
```
```
>>> isinstance(weight_in_pounds('220 lbs'), float)
True
```
```
>>> weight_in_pounds('220 lbs') == 220
True
```
```
>>> weight_in_pounds('196 lbs') == 196
True
```
```
>>> weight_in_pounds('196.5 lbs') == 196.5
True
```


Tests for q5c (points: 2)
```
>>> callable(compute_bmi)
True
```
```
>>> type(compute_bmi('6\' 3"', '220 lbs')) == float
True
```
```
>>> np.isclose(compute_bmi('6\' 3"', '220 lbs'), 27.4951)
True
```
```
>>> np.isclose(compute_bmi('5\' 9"', '171 lbs'), 25.2495)
True
```


Tests for q5d (points: 5)
```
>>> callable(classify_bmi)
True
```
```
>>> isinstance(classify_bmi(30.1689), str)
True
```
```
>>> isinstance(football_roster_bmi, Table)
True
```
```
>>> football_roster_bmi.num_rows == football_roster.num_rows
True
```
```
>>> football_roster_bmi.num_columns == football_roster.num_columns + 2
True
```
```
>>> 'BMI' in football_roster_bmi.labels
True
```
```
>>> 'BMI Category' in football_roster_bmi.labels
True
```
```
>>> classify_bmi(27.4951) == 'Overweight'
True
```
```
>>> np.isclose(football_roster_bmi.where('Name', 'Mac Jones').column('BMI')[0], 27.4951)
True
```
```
>>> football_roster_bmi.where('Name', 'Mac Jones').column('BMI Category')[0] == 'Overweight'
True
```


Tests for q_7ai (points: [0])
```
>>> isinstance(article_i, bool)
True
```


Tests for q_7aii (points: [0])
```
>>> isinstance(article_ii, bool)
True
```


Tests for q_7aiii (points: [0])
```
>>> isinstance(article_iii, bool)
True
```


Tests for q_7b (points: [0, 0])
```
>>> isinstance(bmi_alternatives, np.ndarray)
True
```
```
>>> np.all([x in ['A', 'B', 'C', 'D', 'E', 'F', 'G'] for x in bmi_alternatives])
True
```
